Deciphering deep-sea chemosynthetic symbiosis by single-nucleus RNA-sequencing

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Wang, He et al have constructed comprehensive single nucleus atlas for the gills of the deep sea Bathymodioline mussels, which possess intracellular symbionts that provide a key source of carbon and allow them to live in these extreme environments. They provide annotations of the different cell states within the gills, shedding light on how multiple cell types cooperate to give rise to the emergent functions of the composite tissues and the gills as a whole. They pay special attention to characterizing the bacteriocyte cell populations and identifying sets of genes that may play a role in their interaction with the symbiotes.

Wang, He et al sample mussels from 3 different environments: animals from their native methane rich environment, animals transplanted to a methane-poor environment to induce starvation and animals that have been starved in the methane-poor environment and then moved back to the methane-rich environment. They demonstrated that starvation had the biggest impact on bacteriocyte transcriptomes. They hypothesize that the up-regulation of genes associated with lysosomal digestion leads to the digestion of the intracellular symbiont during starvation, while the non-starved and reacclimated groups more readily harvest the nutrients from symbiotes without destroying them. Further work exploring the differences in symbiote populations between ecological conditions will further elucidate the dynamic relationship between host and symbiote. This will help disentangle specific changes in transcriptomic state that are due to their changing interactions with the symbiotes from changes associated with other environmental factors.

This paper makes available a high quality dataset that is of interest to many disciplines of biology. The unique qualities of this non-model organism and collection of conditions sampled make it of special interest to those studying deep sea adaptation, the impact of environmental perturbation on Bathymodioline mussels populations, and intracellular symbiotes. The authors also use a diverse array of tools to explore and validate their data.

Reviewer #2 (Public Review):

Wang, He et al. shed insight into the molecular mechanisms of deep-sea chemosymbiosis at the single-cell level. They do so by producing a comprehensive cell atlas of the gill of Gigantidas platifrons, a chemosymbiotic mussel that dominates the deep-sea ecosystem. They uncover novel cell types and find that the gene expression of bacteriocytes, the symbiont-hosting cells, supports two hypotheses of host-symbiont interactions: the "farming" pathway, where symbionts are directly digested, and the "milking" pathway, where nutrients released by the symbionts are used by the host. They perform an in situ transplantation experiment in the deep sea and reveal transitional changes in gene expression that support a model where starvation stress induces bacteriocytes to "farm" their symbionts, while recovery leads to the restoration of the "farming" and "milking" pathways.

A major strength of this study includes the successful application of advanced single nucleus techniques to a non-model, deep sea organism that remains challenging to sample. I also applaud the authors for performing an in situ transplantation experiment in a deep sea environment. From gene expression profiles, the authors deftly provide a rich functional description of G. platifrons cell types that is well-contextualized within the unique biology of chemosymbiosis. These findings offer significant insight into the molecular mechanisms of deep-sea host-symbiont ecology, and will serve as a valuable resource for future studies into the striking biology of G. platifrons.

The authors' conclusions are generally well-supported by their results. However, I recognize that the difficulty of obtaining deep-sea specimens may have impacted experimental design and no replicates were sampled.

It is notable that the Fanmao cells were much more sparsely sampled. It appears that fewer cells were sequenced, resulting in the Starvation and Reconstitution conditions having 2-3x more cells after doublet filtering. These discrepancies also are reflected in the proportion of cells that survived QC, suggesting a distinction in quality or approach. However, the authors provide clear and sufficient evidence via bootstrapping that batch effects between the three samples are negligible. While batch effect does not appear to have affected gene expression profiles, the proportion of cell types may remain sensitive to sampling techniques, and thus interpretation of Fig. S12 must be approached with caution.

Reviewer #3 (Public Review):

Wang et al. explored the unique biology of the deep-sea mussel Gigantidas platifrons to understand fundamental principles of animal-symbiont relationships. They used single-nucleus RNA sequencing and validation and visualization of many of the important cellular and molecular players that allow these organisms to survive in the deep-sea. They demonstrate that a diversity of cell types that support the structure and function of the gill including bacteriocytes, specialized epithelial cells that host sulfur-oxidizing or methane-oxidizing symbionts as well as a suite of other cell types including supportive cells, ciliary, and smooth muscle cells. By performing experiments of transplanting mussels from one habitat which is rich in methane to methane-limited environments, the authors showed that starved mussels may consume endosymbionts versus in methane-rich environments upregulated genes involved in glutamate synthesis. These data add to the growing body of literature that organisms control their endosymbionts in response to environmental change.

The conclusions of the data are well supported. The authors adapted a technique that would have been technically impossible in their field environment by preserving the tissue and then performing nuclear isolation after the fact. The use of single-nucleus sequencing opens the possibility of new cellular and molecular biology that is not possible to study in the field. Additionally, the in-situ data (both WISH and FISH) are high-quality and easy to interpret. The use of cell-type-specific markers along with a symbiont-specific probe was effective. Finally, the SEM and TEM were used convincingly for specific purposes in the case of showing the cilia that may support water movement.

The one particular area for future exploration surrounds the concept of a proliferative progenitor population within the gills. The authors recover molecular markers for these putative populations and additional future work will uncover if these are indeed proliferative cells contribute to symbiont colonization.

Overall the significance of this work is identifying the relationship between symbionts and bacteriocytes and how these host bacteriocytes modulate their gene expression in response to environmental change. It will be interesting to see how similar or different these data are across animal phyla. For instance, the work of symbiosis in cnidarians may converge on similar principles of there may be independent ways in which organisms have been able to solve these problems.

We extend our sincere gratitude to all the reviewers for their positive comments and kind words. We highly value the substantial efforts they made in helping us improve and enhance our manuscript. Additionally, we appreciate the reviewers for pointing out the limitations of our current study, which will guide us in improving our future researches.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

This study system is so interesting and this is a truly unique and exciting dataset. Most of my suggestions are aimed at improving readability and making it more accessible for a broader audience, since I predict many fields will find it interesting.

Line 60: which species of mussel? Is this the same one?

We appreciate the comments from the reviewer. The reference here is to deep-sea bathymodiolin mussels, which, in most cases, possess enlarged gill filaments that accommodate symbionts.

Line 237-230: citation of previous findings missing

We appreciate the comments from the reviewer. After carefully reviewing these paragraphs, we believe that all the previous findings have now been properly cited.

Line 256: it might be a good idea to give a brief description of what slingshot analysis is here

We appreciate the comments from the reviewer. We have revise the corresponding part of our manuscript to make it clear.

This parts of manscript now reads: “We performed Slingshot analysis, which uses a cluster-based minimum spanning tree (MST) and a smoothed principal curve to determine the developmental path of cell clusters. The re-sult shows that the PEBZCs might be the origin of all gill epithelial cells, including the other two proliferation cells (VEPC and DEPC) and bacteriocytes (Supplementary Fig. S6).” Line 203-207 of the revised manscript.

Line 289: Wording is a bit confusing- what is meant by morphological analysis?

We acknowledge that our wording might be a bit confusing here. We are referring to the TEM ultrastructural analysis. Therefore, we have changed “morphological analysis” to “ultrastructural analysis.” Line231 in the revised manuscript.

Line 351-354: how did you calculate distances? How many dimensions were used?

We calculated the centroid coordinates for each cell type in each state on the 2-dimensional UMAP plot (Fig. 6A). Then, for each cell type, we determined the Euclidean distance between the centroid coordinates of each pair of states. We have revised the manuscript with this more detailed description. Line 292-295 of revised manuscript.

Line 462: identify -> identified

We apologize for our mistake and appreciate the reviewer’s kind assistance with proofreading. The typo has been corrected in the new version. Line396 of the revised manscript.

Line 509: what does the size of the dot represent?

In this context, the color and intensity of each dot represent a specific gene’s expression level in the single-cell cluster. The dot size is universal and therefore does not convey a specific meaning.

Fig 3A: What is the blue cluster highlighted?

We apologize for our mistake. The label for the teal box was missed. We have corrected our mistake in the revised manuscript.

Fig 3K: Wording in key is confusing.

We have modified our description of Fiugre 3K in the figure legneds. Now it reads: “Schematic of water flow agitated by different ciliary cell types. The color of arrowheads corresponds to water flow potentially influenced by specific types of cilia, as indicated by their color code in Figure 3A.” Line462-464 in the revised manscript.

Fig 5B: which population of mussels was used to take these images?

These mussels from “Fanmao” (methane rich) site were used to take these images. We have revised our material and methods to make it clear. Line602-603 of the revised manuscript.

Fig 5E,5G,5H: panels not referenced in text

We apologize for our mistake and appreciate the reviewer’s thorough reading. This error has been corrected in the new version of the manuscript. Line233 of the revised manuscript.

Reviewer #2 (Recommendations For The Authors):

Minor comments:

Fig. 3A - the teal box in the legend lacks a label

We apologize for our mistake. The label for the teal box was missed. We have corrected our mistake in the

Reviewer #3 (Recommendations For The Authors):

My enthusiasm for the manuscript remains high and I appreciate the authors care in responding to the various reviewer questions and concerns.

In regards to the cell proliferation results, I have modified my public review and look forward to your future work in this area. The data for both pHistone H3 and anti PCNA are compelling!

One typo I did catch occurs on line 520. I believe you meant to say "outer" not "otter."

We apologize for our mistake and appreciate the reviewer’s kind assistance with proofreading. The typo has been corrected in the new version.

Changes from the earlier version to the revised version

Abstract

Abstract

Bathymodioline mussels dominate deep-sea methane seep and hydrothermal vent habitats and obtain nutrients and energy primarily through chemosynthetic endosymbiotic bacteria in the bacteriocytes of their gill. However, the molecular mechanisms that orchestrate mussel host-symbiont interactions remain unclear. Here, we constructed a comprehensive cell atlas of the gill in the mussel Gigantidas platifrons from the South China Sea methane seeps (1100m depth) using single-nucleus RNA sequencing (snRNA-seq) and whole-mount in situ hybridisation. We identified 13 types of cells, including three previously unknown ones, and uncovered unknown tissue heterogeneity. Every cell type has a designated function in supporting the gill’s structure and function, creating an optimal environment for chemosynthesis, and effectively acquiring nutrients from the endosymbiotic bacteria. Analysis of snRNA-seq of in situ transplanted mussels clearly showed the shifts in cell state in response to environmental oscillations. Our findings provide insight into principles of host-symbiont interaction and the bivalves’ environmental adaption mechanisms.

Introduction

Earlier version: Mutualistic interactions between multicellular animals and their microbiota play a fundamental role in animals’ adaptation, ecology, and evolution 1,2. By associating with symbionts, host animals benefit from the metabolic capabilities of their symbionts and gain fitness advantages that allow them to thrive in habitats they could not live in on their own1,3. Prime examples of such symbioses are bathymodioline mussels and gammaproteobacteria endosymbionts4. The bathymodioline mussels occur worldwide at deep-sea chemosynthetic ecosystems, such as cold-seeps, hydrothermal vents, and whale falls4. The host mussel acquired sulfur-oxidising (SOX) and/or methane-oxidising (MOX) symbiont through horizontal transmission at their early life stage3. Their symbionts, which are hosted in a specialised gill epithelial cell, namely, the bacteriocytes5,6, utilise the chemical energy from the reduced chemical compounds such as CH and H2S released from cold seeps or hydrothermal vents to fix carbon and turn into carbon source for the host mussel7,8. The ecological success of the bathymodioline symbioses is apparent: the bathymodiolin mussels are often among the most dominant species in the deep-sea chemosynthetic ecosystems9,10. Thus, it is critical to know how the bathymodioline mussels interact with the symbionts and maintain the stability and efficiency of the symbiosis.

Revised version: Mutualistic interactions between multicellular animals and their microbiota play a fundamental role in animals’ adaptation, ecology, and evolution 1,2. By associating with symbionts, host animals benefit from the metabolic capabilities of their symbionts and gain fitness advantages that allow them to thrive in habitats they could not live in on their own1,3. Prime examples of such symbioses are bathymodioline mussels and gammaproteobacteria endosymbionts4. The bathymodioline mussels occur worldwide at deep-sea chemosynthetic ecosystems, such as cold-seeps, hydrothermal vents, and whale falls4. The host mussel acquired sulfur-oxidising (SOX) and/or methane-oxidising (MOX) symbiont through horizontal transmission at their early life stage3. Their symbionts, which are hosted in a specialised gill epithelial cell, namely, the bacteriocytes5,6, utilise the chemical energy from the reduced chemical compounds such as CH4 and H2S released from cold seeps or hydrothermal vents to fix carbon and turn into carbon source for the host mussel7,8. The ecological success of the bathymodioline symbioses is apparent: the bathymodiolin mussels are often among the most dominant species in the deep-sea chemosynthetic ecosystems9,10. Thus, it is critical to know how the bathymodioline mussels interact with the symbionts and maintain the stability and efficiency of the symbiosis.

The gill structure of bathymodiolin mussels has undergone remarkable adaptations at the molecular, cellular, and tissue levels to support their deep-sea chemosynthetic lifestyle11,12. Compared to shallow water mussels13, the gill filaments of bathymodiolin mussels have enlarged surfaces allowing them to hold more symbionts per filament (Supplementary Fig. S1). This adaption requires not only novel cellular and molecular mechanisms to maintain and support the enlarged gill filament structure, but also a strong ciliary ventilation system to pump vent/seep fluid to fuel the symbiotic bacteria. Previous studies based on whole-genome sequencing and bulk RNA-seq projects have shown that genes in the categories of nutrient transporters, lysosomal proteins, and immune receptors are either expanded in the host mussel’s genome or up-regulated in the gill12, 14–17, suggesting that these genes involved in host-symbiont interaction. Though providing deep molecular insights, these studies mainly used homogenised tissues that average genes’ expression levels amongst different cell types and eliminate cell and gene spatial distribution information18–21. In addition, the broad expression and function of potentially ‘symbiosis-related’ proteins also greatly limited data interpretation. Therefore, a systemic atlas of gill cell types and the descriptions of cell-type-specific gene expressional profiles are warranted to a better understanding of the host-symbiont interaction and environmental adaptation mechanisms of the bathymodioline symbiosis.

In recent years, single cell/nucleus RNA-sequencing (sc/sn RNA-seq) technologies have become one of the preferred methods for investigating the composition of complex tissue at the transcriptional level22. snRNA-seq has several advantages, such as compatibility with frozen samples, elimination of dissociation-induced transcriptional stress responses, and reduced dissociation bias23. These advantages are significant for deep-sea and other ecological studies because cell and molecular biology facilities are not available in the field. To examine the cellular and molecular mechanisms of the environmental adaptations and host-symbiont interactions in bathymodioline mussels, we conducted a single nucleus RNA-sequencing-based transcriptomic study. We analysed the gill symbiosis in the dominate deep-sea mussels inhabiting the F-site cold-seep Gigantidas platifrons24, which hosts a single a MOX endosymbiont population that consist of only one 16S rRNA phylotype. Our work provides a proof-of-principle for environmental adaptation mechanisms study in non-traditional but ecologically important organisms with single-nucleus RNA-sequencing technologies.

Results and discussion

G. platifrons deep-sea in situ transplant experiment and single-cell transcriptomic sequencing

We conducted a G. platifrons deep-sea in situ transplant experiment at the “F-Site” Cold-seep (∼1117 m depth) and retrieved three groups of samples (Fig. 1A), as follows: the ‘Fanmao’ (meaning prosperous) group, which comprised the mussels collected from methane-rich (∼40 μM) site of the cold-seep, where the animals thrived; the ‘starvation’ group, which comprised the mussels collected from the methane-rich site and then moved to a methane-limited (∼0.054 μM) starvation site for 14 days before retrieval; and the ‘reconstitution’ group, which comprised the methane-rich site mussels moved to the ‘Starvation’ site for 11 days and then moved back to the methane-rich site for another 3 days.

Earlier version: Identification of 14 cell types in the gill of deep-sea symbiotic mussel Gigantidas platifrons. A: Overall experimental scheme describing the deep-sea in situ transplant experiment, the sample preparation procedures and the single-cell analysis and validation pipeline. Three G. platifrons samples were included in the present study: ‘Fanmao,’ Starvation and Reconstitution. The cell nucleus was extracted from each sample, which included a pool of gill posterior tip of three mussels. The snRNA-seq libraries were constructed according to the BD Rhapsody singlenuclei 3′ protocol. Cell population-specific markers were validated by WISH and ISH. B: The image shows the posterior end of the gill of G. platifrons. C: UMAP representation of G. platifrons gill single cells. Cell clusters are coloured and distinctively labelled. D: Heat map profile of markers in each cluster. The colour gradient represents the expression level of every single cell.

Revised version: Identification of 14 cell types in the gill of deep-sea symbiotic mussel Gigantidas platifrons. A: Overall experimental scheme describing the deep-sea in situ transplant experiment, the sample preparation procedures and the single-cell analysis and validation pipeline. Three G. platifrons samples were included in the present study: ‘Fanmao,’ Starvation and Reconstitution. The cell nucleus was extracted from each sample, which included a pool of gill posterior tip of three mussels. The snRNA-seq libraries were constructed according to the BD Rhapsody single-nuclei 3′ protocol. Cell population-specific markers were validated by WISH and ISH. B: The image shows the posterior end of the gill of G. platifrons. C: UMAP representation of G. platifrons gill single cells. Cell clusters are coloured and distinctively labelled. D: Heat map profile of markers in each cluster. The colour gradient represents the expression level of every single cell.

We next performed single nucleus RNA-sequencing (snRNA-seq) of the gill posterior tip of the three groups of samples using the microwell-based BD Rhapsody platform™ Single-Cell Analysis System (Fig. 1A). After quality control, we obtained 9717 (Fanmao), 21614 (Starvation), and 28928 (Reconstitution) high-quality single nuclei transcriptomic data.

Cell atlas of G. platifrons gill

To unravel the intricate cellular composition, we utilised a reciprocal PCA (RPCA) strategy to project cells in the three samples onto the same space based on conserved expressed genes among them25. This strategy maximises the number of cells per cluster regardless of the organism’s state and therefore maximises genes per cluster26.

Through the implementation of Seurat, we revealed 14 cell clusters, each was associated with a set of marker genes (Figs. 1B-D) 27,28. Given the limited availability of canonical marker genes for G. platifrons and molluscs in general, we undertook a meticulous approach to characterize each cell cluster. This involved: 1) examining the cluster marker genes’ functions, 2) identifying the expression pattern of cluster marker gene using whole-mount in situ hybridisation (WISH) and/or double-fluorescent in situ hybridisation (FISH) analyses, and 3) conducting scanning electron microscopy (SEM) and transmission electron microscopy (TEM) analyses. We successfully identified and characterized 13 of the 14 cell clusters, which could be categorized into four major groups, including 1) the supportive cells, 2) the ciliary cells, 3) the proliferation cells, and 4) the bacteriocytes. One cell cluster remained ambiguous. To assess the robustness of each cell cluster, we employed a bootstrap sampling and clustering algorithm examining the similarity among clusters and obtained strong support for all clusters through a combined analysis of the three samples (Supplementary Fig. S2A). Furthermore, when examining each sample individually, we found that the majority of clusters demonstrated robust support, with the exception of the three ciliary cell clusters, which showed overlaps of assignment probabilities among them (Supplementary Fig. S2 B-D). These three cell types are derived from a same precursor, and exhibited relatively lower numbers of nuclei, resulting in a reduced availability of genes per cell type, which is particularly true for the food grove ciliary cell (Supplementary Table S1). By integrating the spatial and functional annotations of these cell clusters, we gained insight into their collective efforts in supporting the symbiotic relationship and maintaining the high-efficiency chemosynthetic system within the gill tissue.

Supportive cells

We have identified four supportive cell types, namely, inter lamina cells basal membrane cells (BMC)1, BMC2, and mucus cells (Figs 2A, B and C).

Earlier version: Supportive cell populations of G. platifrons gill. A: UMAP representation of the four supportive cell populations. B: Expression profiles of the cell markers that are specific or enriched in the supportive cell populations. The sizes of the circles represent the percentages of cells in those clusters that expressed a specific gene. Genes shown in red were validated by WISH or double FISH. C and C’: Schematics demonstrating the overall structural (panel C) and supportive cell distribution (C’). D, F, H, and J: WISH characterisation of the selected representative cell population markers. E, G, I, and K: Double FISH characterisation of the selected representative cell population markers. The white arrowheads in G indicate the BMC1 cells locates at the otter rim of gill slice. Scale bar: 50µm.

Revised version: Supportive cell populations of G. platifrons gill. A: UMAP representation of the four supportive cell populations. B: Expression profiles of the cell markers that are specific or enriched in the supportive cell populations. The sizes of the circles represent the percentages of cells in those clusters that expressed a specific gene. Genes shown in red were validated by WISH or double FISH. C and C’: Schematics demonstrating the overall structural (panel C) and supportive cell distribution (C’). D, F, H, and J: WISH characterisation of the selected representative cell population markers. E, G, I, and K: Double FISH characterisation of the selected representative cell population markers. The white arrowheads in G indicate the BMC1 cells locates at the outer rim of gill slice. Scale bar: 50µm.

The inter lamina cells, marked by high expression of fibrillar-forming collagens, were the cells located between the two layers of basal membranes (Fig 2D and E). These cells were previously identified as amoeboid cells13, and their function has not been explored. The inter lamina cells were densely distributed around the food groove and at the rim of the gill filament. While on the middle part of the gill filament, the inter lamina cells distributed in parallel rows. It might help to connect the two sheets of basal membranes and maintain the spatial integrity of the gill filament. In addition, we observed the enrichment of ribosomal proteins involved in protein synthesis (Supplementary Table S2), indicating a high metabolic rate29 of inter lamina cells. This implies that the inter lamina cells may process the nutrients acquired from the symbiont.

We have also identified two types of BMC populations which have not been recognised before. The BMC1 (Fig. 2F and G) expressed genes encoding extracellular matrix and adhesive proteins (Fig. 2B and Supplementary Table S2), such as SCO-spondin (Bpl_scaf_27638-4.18), Zonadhesin (Bpl_scaf_1371-2.4), Cartilage matrix proteins (Bpl_scaf_16371-0.14 and Bpl_scaf_3387-5.47), and Collagen alpha-1(XXI) chain protein (Bpl_scaf_51350-0.19). BMC2 (Fig. 2H and I) was distinctive from BMC1 in the context of expression (Fig 2A) and marker genes (Fig 2B). Genes encoding hemicentins (Bpl_scaf_38731-0.50 and Bpl_scaf_7293-0.14), which are known to stabilise the cells’ contact with the basal membrane30, were highly expressed in BMC2. Thus, these two cell types are likely to help building the basal lamina, and stabilising the epithelial-derived cells, such as bacteriocytes and intercalary cells, on the surface of the basal membrane.

Mucus cells are specialised secretory cells with intracellular mucus vacuoles13. In shallow-water filter-feeding bivalves, mucus cells secret mucus, which cooperates with the ciliary ventilation system to capture, process, and transport food particles to their mouths, known as filter feeding mechanism31–34. Deep-sea mussels may not necessarily retain this function since they do not normally acquire food resources such as phytoplankton and planktonic bacteria but obtain most of nutrient through their chemosynthetic symbiont. Thus, it was hypothesised that mucus cells are involved in other biological functions such as immune responses to pathogens18. Herein, genes encoding proteins with microbe-binding functionalities were enriched in mucus cells (Fig. 2B and Supplementary Table S2), such as C-type lectins (Bpl_scaf_58694-0.22 and Bpl_scaf_38133-0.8), C1q domain-containing protein (Bpl_scaf_55216-2.11), and Fibrinogen C domain-containing protein (Bpl_scaf_18519-2.15)35,36. The expression of similar immune genes was upregulated in a shallow water mussel Mytilus galloprovincialis when challenged by a pathogenic bacteria37. Our WISH and double-FISH analyses showed that mucus cells were embedded within the outer rim cilia and scattered on the gill lamella alongside the bacteriocytes (Figs. 2J and 2K). These data collectively suggest that mucus cells may help mussel maintain the immune homeostasis of gill. Interestingly, WISH and double-FISH analyses showed that mucus cells were also distributed alongside the gill lamella’s food groove and the inner edge, where the density of bacteriocytes is low (Fig. 2J). Because the food groove is the main entry of food practical to the labial palps and mouth38, this distribution pattern implies that mucus cells may be also involved in capturing planktonic bacteria and sending them to the mouth.

Ciliary and smooth muscle cells

A remarkable feature of bathymodioline mussel’s gill is its ciliary ventilation system, which constantly agitates the water and provides the symbiont with the necessary gas39. We identified four types of ciliary cells (Figs. 3A and 3B): apical ridge ciliary cells (ARCCs), food groove ciliary cells (FGCCs), lateral ciliary cells (LCCs), and intercalary cells (ICs), as well as a newly identified type of smooth muscle cells. All ciliary cells were marked by canonical cilium genes, such as genes encoding flagellar proteins, ciliary motor kinesin proteins, ciliary dynein proteins, and ciliary microtubules and were clearly distinguishable by specifically expressed genes (Supplementary Table S3 and Supplementary Fig. S3). ARCCs were characterized by expression of Tubulin alpha-1A chain Bpl_3489-0.37 (Fig. 3C) and Bpl_scaf_20631-1.16, which encodes homeobox Dlx6a-like protein (Fig. 3D) which is a marker of the apical ectodermal ridge40. FGCCs, which could be ladled by expression of marker gene Bpl_scaf_5544-0.0 (Fig. 3E), expressed genes encoding primary cilia development regulator Tubby-related protein 341 (Bpl_scaf_24834-2.3) and primary ciliary cell structural protein TOG array regulator of axonemal microtubules protein 142,43 (Bpl_scaf_55620-1.3), suggesting the FGCCs are the sensory ciliary cells that gather information from the surrounding environment. Both ARCCs and FGCCs were located around the ventral tip of the gill filament (Fig. 3F). LCCs were distributed as two parallel rows along the gill lamella’s outer rim and ciliary disks’ outer rim (Figs. 3G and 3H) and had highly expressed genes involved in cilium structure, cilium movement and ATP synthases (Fig. 3B and Supplementary Table S3), indicating that LCCs may have a strong ability to beat their cilia. Interestingly, we identified a group of previously unreported smooth muscle cells (SMCs) co-localised with the LCCs as showed by WISH (Fig. 3I and Supplementary Fig. S4). SMCs strongly expressed several Low-density lipoprotein receptors (LDL receptor) and LDLR-associated proteins44 (Fig. 3B). In addition, these cells also expressed the angiotensin-converting enzyme-like protein45 (Bpl_scaf-29477-5.10) and the ‘molecular spring’ titin-like protein46 (Bpl_scaf_56354-7.8) (Supplementary Table S3). The expression of these genes could be commonly found in human vascular smooth muscle cells47–49. Collectively, we suspect that SMCs are involved in lateral cilium movement and the gill slice contraction.

Ciliary cell populations of G. platifrons gill. A: UMAP representation of the four ciliary cell populations and potential smooth muscle cell population. B: Expression profiles of the cell markers that are specific or enriched in the ciliary cell populations. The sizes of the circles represent the percentages of cells in those clusters that expressed a specific gene. The genes shown in red were validated by WISH or double FISH. C-E, G and I-J: WISH characterisation of the selected representative cell population markers. F and H: SEM analysis of the ciliary cells of G. platifrons gill. K: Schematic of the water flow agitated by different ciliary cell types. The coluor of arrowheads corresponds to water flow potentially influenced by specific types of cilia, as indicated by their colour code in Figure 3A.

ICs are the specialised ciliary cells surrounding the bacteriocytes (Figs. 3B and 3J). WISH analysis showed that the expressions of IC markers had apparent spatial variations (Fig. 3J). The expression of ICs marker rootletin50–52 (Bpl_scarf_22362-6.9) was considerably higher in the ICs close to the frontal edge of the gill filament, and the expression gradually decreased along with the direction of inter lamina water flow (Fig. 3K and Supplementary Fig. S5), implying that ICs ventilate the water flow and the mucus through the gill filaments. Furthermore, compared with the other three types of ciliary cells, the ICs expressed several genes encoding transcription factors involved in determining cell fate (Fig. 3B and Supplementary Table S3), such as transcription factors Sox 8 (Bpl_scarf_40595-5.30) and Wnt pathway cell polarity regulator secreted frizzled-related protein 5 (Bpl_scaf_57424-6.1)53, suggesting that the ICs might also play regulatory roles54–56.

Proliferation cells

It has long been known that the bathymodioline mussel gill has three types of proliferation cells that are conserved throughout all filibranchia bivalves: the budding zone at the posterior end of the gill where new gill filaments are continuously formed, and the dorsal and ventral ends of each gill filaments57,58. These “cambial-zone”-like cell populations could continuously proliferate throughout the whole life span of the mussel58. Our snRNA-seq data recognised three types of proliferation cells, which is consistent with previous findings (Fig. 4A, Fig. 4B and Supplementary Table S4). The gill posterior end budding zone cells (PEBZCs) are located on the first few freshly proliferated filaments of the posterior tip of gill (Figs. 4C and 4D). The PBEZCs marker (Bpl_scaf_61993-0.4) gradually disappeared at around the 11th-12th row of the gill filament (Fig. 4D), suggesting the maturation of the gill filaments, which is similar to the developmental pattern reported in another deep-sea mussel Bathymodiolus azoricus 59. The dorsal end proliferation cells (DEPCs), which expressed the hallmarks genes of muscular tissue (Fig. 4B), as well as cell proliferation and differentiation regulators, were the proliferation cells in connective tissue at the dorsal end of the gill slice (Fig. 4A, 4E). The ventral end proliferation cells (VEPCs) were two symmetrical triangle-shaped cell clusters of small symbiont-free cells (Fig. 4F). In VEPCs, genes encoding ribosomal proteins, chromatin proteins, RNA and DNA binding proteins, and cell proliferation markers were all up-regulated (Fig. 4B and Supplementary Table S4), indicating that VEPCs are meristem-like cells59,60.

Proliferation cell populations of G. platifrons gill. A: UMAP representation of the three proliferation cell populations. B: Expression profiles of the cell markers that are specific or enriched in the supportive cell populations. The sizes of the circles represent the percentages of cells in those clusters that expressed a specific gene. Genes shown in red were validated by WISH or double FISH. C and D: Photographic and schematic analyses of the spatial position of the three Proliferation cell populations. E, F and J: FISH and WISH characterisation of the selected population markers. The marker genes confirmed by ISH or WISH in the current study are indicated in red. Scale bar: 50 µm.

Removed in the revised version: It has been hypothesised that new proliferation cells will be colonised by symbionts, serving as the vital mechanism of bacteriocyte recruitment59. We then determined which proliferation cell type gives rise to the mature cells, especially the bacteriocytes, which has little available information regarding their precursor and transmission mode in bathymodioline mussels 59,61.

Earlier version: Slingshot trajectory analysis suggested that the PEBZCs might be the origin of all gill epithelial cells, including the other two proliferation cells (VEPC and DEPC) and bacteriocytes (Supplementary Fig. S6). The sole exception was BMC2, which may be derived from VEPC rather than PEBZC. This result is consistent with previous studies which suggested new gill filaments of the filibranch mussels are formed in the gill’s posterior budding zones59. The colonisation by the symbiont might play a crucial role in determining the fate of the bacteriocytes. Noticeably, in G. platifrons, only the pillar-shaped first row of gill filament, comprised of small meristem-like cells, was symbiont-free (Figs. 4D and D’), whereas all the other gill filaments were colonised by symbionts. This pattern of symbiosis establishment is different from that of B. azoricus, in which PEBZCs are symbiont free and are gradually colonised by symbionts released from the bacteriocytes on the adjacent mature gill filaments after maturation59. On mature gill filaments, the DEPCs and VEPCs are seemingly the sources of new cells that sustain the growth of the gill filament from both dorsal and ventral directions, respectively57. Interestingly, comparable active ventral and dorsal end proliferation zones have also been identified in the symbiotic mussel B. azoricus, whereas they are absent in the shallow-water mussel Mytilus edulis62. This contrast further suggests the potential involvement of DEPCs and VEPCs in the establishment of symbiosis.

Revised version: It has been hypothesised that new proliferation cells will be colonised by symbionts, serving as the vital mechanism of bacteriocyte recruitment59. We then determined which proliferation cell type gives rise to the mature cells, especially the bacteriocytes, which has little available information regarding their precursor and transmission mode in bathymodioline mussels 59,61. We performed Slingshot analysis, which uses a cluster-based minimum spanning tree (MST) and a smoothed principal curve to determine the developmental path of cell clusters. The result shows that the PEBZCs might be the origin of all gill epithelial cells, including the other two proliferation cells (VEPC and DEPC) and bacteriocytes (Supplementary Fig. S6). The sole exception was BMC2, which may be derived from VEPC rather than PEBZC. This result is consistent with previous studies which suggested new gill filaments of the filibranch mussels are formed in the gill’s posterior budding zones59. The colonisation by the symbiont might play a crucial role in determining the fate of the bacteriocytes. Noticeably, in G. platifrons, only the pillar-shaped first row of gill filament, comprised of small meristem-like cells, was symbiont-free (Figs. 4D and D’), whereas all the other gill filaments were colonised by symbionts. This pattern of symbiosis establishment is different from that of B. azoricus, in which PEBZCs are symbiont free and are gradually colonised by symbionts released from the bacteriocytes on the adjacent mature gill filaments after maturation59. On mature gill filaments, the DEPCs and VEPCs are seemingly the sources of new cells that sustain the growth of the gill filament from both dorsal and ventral directions, respectively57. Interestingly, comparable active ventral and dorsal end proliferation zones have also been identified in the symbiotic mussel B. azoricus, whereas they are absent in the shallow-water mussel Mytilus edulis62. This contrast further suggests the potential involvement of DEPCs and VEPCs in the establishment of symbiosis.

Bacteriocytes and host-symbiont interaction

We conducted the whole-mount FISH using a bacterial 16S rRNA probe for symbiont to determine the spatial distribution of bacteriocytes, and their positions relative to the other cell types on the gill filament (Figs. 5A, B). Bacteriocytes covered the majority of the surface of the gill filament, except the ventral tip, the ciliary disk, and the frontal edge (lateral ciliary). The bacteriocytes were surrounded by intercalary cells with microvilli and cilia on the surface (Fig. 5C).

Characterisation of the bacteriocytes of G. platifrons. A: Schematic of the overall structure of G. platifrons gill filaments. B: Whole-mount FISH analyses of the overall distribution of bacteriocytes on the G. platifrons gill filament. C: SEM analysis of the bacteriocytes. D: TEM analysis of a bacteriocyte. E: UMAP representation of G. platifrons bacteriocytes. F: Expression profiles of the cell markers that are specific or enriched in the bacteriocytes. The sizes of the circles represent the percentages of cells in those clusters that expressed a specific gene. Double FISH validated the genes shown in red. I: Schematic of the host-symbiont interaction based on the single-cell transcriptome of G. platifrons bacteriocytes. The marker genes confirmed by ISH in the current study are indicated in red. Scale bar in panel G and H: 25 µm.

Earlier version: The gene expression profile of bacteriocytes aligned well with the morphological analysis, which suggested that the bacteriocytes have structural, metabolic, and regulatory adaptions to cope with the symbiont (Figs. 5D, 5F and 5I). It has been hypothesised that the bacteriocytes extract nutrients from the symbiont through two pathways: lysing the symbiont through the lysosome (the ‘farming’ pathway) or directly utilising the nutrient produced by the symbiont (the ‘milking’ pathway)63,64. Our TEM observations clearly detected intracellular vacuole and lysosome system of the bacteriocyte that harbour, transport, and digest the symbionts (Fig. 5D), which is consistent with previous studies18,65. The snRNA data showed bacteriocytes expressed cellular membrane synthesis enzyme (phosphoethanolamine N-methyltransferase, Bpl_scaf_15282-3.10) and a series of lysosomal proteins such as lysosomal proteases (cathepsins; Bpl_scaf_61711-5.12, Bpl_scaf_46838-6.40 and Bpl_scaf_59648-4.5; lysosomal alpha-glucosidase-like), protease regulators 56 (TMPRSS15 proteins; Bpl_scaf_52188-1.15 and Bpl_scaf_15410-0.9), and lysosomal traffic regulator proteins (rabenosyn-5, Bpl_scaf_54816-0.3 and Bpl_scaf_52809-1.6, lysosomal-trafficking regulator-like isoform X1). Among the proteases, cathepsins were thought to be evolutionary conserved molecular tools that host utilized to control the residence of their symbiont microbes66. They were also highly expressed in symbiotic tissue of other deep-sea chemosynthetic animals, such as vesicomyid clams and vestimen-tiferan tubeworms67–69. Bacteriocytes also expressed genes encoding cellular vesicle transports (kinesins, Bpl_scaf_14819-0.11, Bpl_scaf_54265-1.13, and Bpl_scaf_4784-1.40)53, potential amino acid transporter57 (Bpl_scaf_36159-5.5), and genes involved in intracellular vesicle transport, such as the FYVE and coiled-coil domain-containing protein 158 (Bpl_scaf_33726-5.7). In addition to form and mobilise early endosomes, the protein products of these genes could transport symbiont-secreted nutrient vesicles to the host (Fig. 5F and Supplementary Table S5), supporting the “milking” pathway.

Revised version: The gene expression profile of bacteriocytes aligned well with the ultrastructural analysis, which suggested that the bacteriocytes have structural, metabolic, and regulatory adaptions to cope with the symbiont (Figs. 5D-5I). It has been hypothesised that the bacteriocytes extract nutrients from the symbiont through two pathways: lysing the symbiont through the lysosome (the ‘farming’ pathway) or directly utilising the nutrient produced by the symbiont (the ‘milking’ pathway)63,64. Our TEM observations clearly detected intracellular vacuole and lysosome system of the bacteriocyte that harbour, transport, and digest the symbionts (Fig. 5D), which is consistent with previous studies18,65. The snRNA data showed bacteriocytes expressed cellular membrane synthesis enzyme (phosphoethanolamine N-methyltransferase, Bpl_scaf_15282-3.10) and a series of lysosomal proteins such as lysosomal proteases (cathepsins; Bpl_scaf_61711-5.12, Bpl_scaf_46838-6.40 and Bpl_scaf_59648-4.5; lysosomal alpha-glucosidase-like), protease regulators 56 (TMPRSS15 proteins; Bpl_scaf_52188-1.15 and Bpl_scaf_15410-0.9), and lysosomal traffic regulator proteins (rabenosyn-5, Bpl_scaf_54816-0.3 and Bpl_scaf_52809-1.6, lysosomal-trafficking regulator-like isoform X1). Among the proteases, cathepsins were thought to be evolutionary conserved molecular tools that host utilized to control the residence of their symbiont microbes66. They were also highly expressed in symbiotic tissue of other deep-sea chemosynthetic animals, such as vesicomyid clams and vestimentiferan tubeworms67–69. Bacteriocytes also expressed genes encoding cellular vesicle transports (kinesins, Bpl_scaf_14819-0.11, Bpl_scaf_54265-1.13, and Bpl_scaf_4784-1.40)53, potential amino acid transporter57 (Bpl_scaf_36159-5.5), and genes involved in intracellular vesicle transport, such as the FYVE and coiled-coil domain-containing protein 158 (Bpl_scaf_33726-5.7). In addition to form and mobilise early endosomes, the protein products of these genes could transport symbiont-secreted nutrient vesicles to the host (Fig. 5F and Supplementary Table S5), supporting the “milking” pathway.

The symbiont of G. platifrons belongs to type I methanotrophy, of which the core metabolic function is linked with the development of intracytoplasmic membranes leading to a high lipid/biomass content70,71. Recent lipid biomarker analyses showed that the gill of G. platifrons contains a high amount of bacterial lipids, which are directly utilised by the host to synthesise most of its lipid contents67. Downstream of the bacteriocyte’s metabolic cascade, genes encoding proteins that may be involved in fatty acid/lipid metabolism, such as perilipin2 (Bpl_scaf_27158-3.8), which is the critical protein to form intracellular lipid droplets72, and a variety of fatty acid metabolism enzymes (acetyl carboxylase 2, Bpl_scaf_55250-5.1173; fatty acid desaturase 1-like isoform X1, Bpl_scaf_35916-0.674,75; long-chain fatty acid-ligase ACSBG2-like isoform X1, Bpl_scaf_28862-1.5 76,77), were up-regulated, suggesting that the fatty acid could be a major form of nutrients passing from the symbiont to the host mussel.

Earlier version: Additionally, bacteriocytes expressed several solute carriers, including sodium/ascorbate cotransporter (Fig. 5I, solute carrier family 23 members 1-like, Bpl_scaf_32311-1.19), sodium/potassium/calcium exchanger (sodium potassium calcium exchanger 3-like, Bpl_scaf_14503-1.28), sodium/chloride ion cotransporter (solute carrier family 12 member 3-like isoform X1, Bpl_scaf_44604-3.7), ferrous iron transporter (solute carrier family 40 member 1-like, Bpl_scaf_63447-0.12) and zinc transporter (zinc transporter ZIP14-like, Bpl_scaf_44428-5.3). The solute carriers are a large family of ATP-dependent transporters that shuttle a variety of small molecules across the cellular membrane78. In several model symbiotic systems, solute carriers play a vital role in host-symbiont interaction by either providing the symbiont with sub-strates79,80 or transporting symbiont-produced nutrients to the host 81–83. Previous studies demonstrated that solute carrier genes were expanded in deep-sea chemosynthetic animals’ genome84 and highly expressed in symbiotic organs85 including bathymodioline mussel’s gill12. In bathymodioline mussels, previous bulk RNA-seq studies detected up-regulated expression of a large variety of solute carriers in the gill, which is consistent with the present study, suggesting that solute carriers may play crucial roles in shuttling nutrients in and out of bacteriocytes and in maintaining the suitable intracellular micro-environment (such as the SLC23A1 and SLC39A14) for the symbiont86,87.

Revised version: Additionally, bacteriocytes expressed several solute carriers, including sodium/ascorbate co-transporter (Fig. 5I, solute carrier family 23 members 1-like, Bpl_scaf_32311-1.19), sodium/potassium/calcium exchanger (sodium potassium calcium exchanger 3-like, Bpl_scaf_14503-1.28), sodium/chloride ion cotransporter (solute carrier family 12 member 3-like isoform X1, Bpl_scaf_44604-3.7), ferrous iron transporter (solute carrier family 40 member 1-like, Bpl_scaf_63447-0.12) and zinc transporter (zinc transporter ZIP14-like, Bpl_scaf_44428-5.3). The solute carriers are a large family of ATP-dependent transporters that shuttle a variety of small molecules across the cellular membrane78. In several model symbiotic systems, solute carriers play a vital role in host-symbiont interaction by either providing the symbiont with substrates79,80 or transporting symbiont-produced nutrients to the host 81–83. Previous studies demonstrated that solute carrier genes were expanded in deep-sea chemosynthetic animals’ genome84 and highly expressed in symbiotic organs85 including bathymodioline mussel’s gill12. In bathymodioline mussels, previous bulk RNA-seq studies detected up-regulated expression of a large variety of solute carriers in the gill, which is consistent with the present study, suggesting that solute carriers may play crucial roles in shuttling nutrients in and out of bacteriocytes and in maintaining the suitable intracellular micro-environment (such as the SLC23A1 and SLC39A14) for the symbiont86,87.

Cell-type-specific response to environmental stresses

Earlier version: To examine cell-type-specific acclimatisation to environmental changes the expression profiles of differentially expressed genes (DEGs) were compared between the three states (Fanmao, starved and reconstituted animals) of samples collected in our in situ transplantation experiment (full lists of DEGs are shown in Supplementary Table S6-S18). Notably, for each state three mussels were processed but pooled for nuclei extraction before snRNA sequencing (see Method for detail). We tested our hypothesis here that at the starvation site, a relatively low concentration of methane shall upset symbiont metabolism and thus substantially affect symbiont-hosting bacteriocytes by assessing the transcriptional changes per cell type. We calculated the distances between the three sets of samples on the UMAP plot (Fig. 6A) estimating the centroid coordinates per sample per cell type, and calculating the Euclid distance between the centroid coordinates of each pair of states per cell type (Supplementary Table S19). The impact of starvation was variable across cell types, as reflected by the cross-state distances (Fig. 6B, green bar). Starvation resulted in the most significant transcriptional changes in bacteriocytes reflected by a large Fanmao-vs-starvation distance (2.3), followed by VEPC and inter lamina cells (2.1 and 1.4, respectively). On the other hand, starvation had little impact on the transcriptions of ciliary cells and most supportive cells such as the food grove ciliary cell, BMC2, and mucus cells (distances <0.5). Fig 6B also shows that after reconstitution (moving the mussels back to the methane-rich site Fanmao for 3 days), the expressional profile of bacteriocytes rapidly changed back, reflected by a large starvation-vs-reconstitution distance and much smaller Fanmao-vs-reconstitution distance (2.3 vs. 0.8; Suppl. Tab. S19). This result coincided with and was supported by our pseudo-time analysis for bacteriocytes (Fig. 6C), showing that bacteriocytes in the reconstitution are in intermediate and transitional states between Fanmao and starvation.

Revised version: To examine cell-type-specific acclimatisation to environmental changes the expression profiles of differentially expressed genes (DEGs) were compared between the three states (Fanmao, starved and reconstituted animals) of samples collected in our in situ transplantation experiment (full lists of DEGs are shown in Supplementary Table S6-S18). Notably, for each state three mussels were processed but pooled for nuclei extraction before snRNA sequencing (see Method for detail). We tested our hypothesis here that at the starvation site, a relatively low concentration of methane shall upset symbiont metabolism and thus substantially affect symbiont-hosting bacteriocytes by assessing the transcriptional changes per cell type. We calculated the centroid coordinates for each cell type in each state on the 2-dimensional UMAP plot (Fig. 6A). Then, for each cell type, we determined the Euclidean distance between the centroid coordinates of each pair of states (Supplementary Table S19). The impact of starvation was variable across cell types, as reflected by the cross-state distances (Fig. 6B, green bar). Starvation resulted in the most significant transcriptional changes in bacteriocytes reflected by a large Fanmao-vs-starvation distance (2.3), followed by VEPC and inter lamina cells (2.1 and 1.4, respectively). On the other hand, starvation had little impact on the transcriptions of ciliary cells and most supportive cells such as the food grove ciliary cell, BMC2, and mucus cells (distances <0.5). Fig 6B also shows that after reconstitution (moving the mussels back to the methane-rich site Fanmao for 3 days), the expressional profile of bacteriocytes rapidly changed back, reflected by a large starvation-vs-reconstitution distance and much smaller Fanmao-vs- reconstitution distance (2.3 vs. 0.8; Suppl. Tab. S19). This result coincided with and was supported by our pseudo-time analysis for bacteriocytes (Fig. 6C), showing that bacteriocytes in the reconstitution are in intermediate and transitional states between Fanmao and starvation.

Analysis of cell population-specific DEGs. A: UMAP representation of the impact of deep-sea in situ transplant treatments on the gene expression pattern of each cell population. The cells from different treatments were labelled with different colours. The dashed line encircled bacteriocyte populations have a considerably altered expression profile. B: Histogram of cross-state distances between the centroids of the Fanmao, Starvation and Reconstitution groups per cell type on UMAP. The black dashed lines indicate the bacteriocyte populations whose expression profile was remarkably altered. C: Visualization of bacteriocytes onto the pseudo time map using monocle. The black lines indicate the main path of the pseudotime ordering of the cells. D: Bifurcation of selected gene expression along two branches in response to environmental perturbation. Genes are clustered hierarchically into two groups, illustrating up- (cluster 1) and down- (cluster2) regulated genes in the starvation state compared with Fanmao. Genes in red colour were discussed in Section 4. The heat map showing the gene expression profiles of all bacteriocytes’ DEGs is shown in supplementary Fig. S7; E: Proposed model for the molecular mechanisms of host-symbiont interactions in response to environmental changes.

For the bacteriocytes population, we conducted cell trajectory pseudo-timing and detailed DEG analysis to interpret the mechanism of host-symbiont interactions. It is important to acknowledge that our sampling strategy might have limitations, as more closely spaced time points could enhance the confidence of trajectory reconstruction. The branched heatmap showed both up- and down-regulated genes during starvation and reconstitution compared with the Fanmao state (Fig. 6D and supplementary Fig. S7). The Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway enrichment analysis of the bacteriocytes’ DEGs provides an overall view of the pathways enriched in each environmental state (Supplementary Fig. S8 A-C). The genes encoding ribosomal proteins are highly expressed in the bacteriocytes under the methane-rich ‘Fanmao’ state (Supplementary Fig. S8A), suggesting an active protein synthesis and cellular metabolism 88. Moreover, organic ion transporters (A BCB P-glycol, Bpl_scaf_18613-16.10; canalicular multispecific organic anion transporter 2, Bpl_scaf_52110-4.43; sodium- and chloride-dependent glycine transporter 1-like, Bpl_scaf_13642-5.8), which may be involved in transporting symbiont-produced nutrients, are also highly expressed (Supplementary Fig. S9).

Earlier version: In starved G. platifrons, negative regulators of cell proliferation, such as autocrine proliferation repressors and receptor-type tyrosine phosphatase beta, were up-regulated in bacterio-cytes, suggesting repression of cell growth. We also observed the enrichment of genes in the apoptosis pathway (Supplementary Fig. S8B), which is attributed to the upregulation of caspases (Bpl_scaf_25165-3.8 and Bpl_scaf_24225-0.12) and cathepsins (Bpl_scaf_64706-0.16, Bpl_scaf_61711-5.12) which can trigger caspases-dependent cell death, suggesting cellular stress condition. Correspondingly, a gene encoding baculoviral IAP repeat-containing protein (Bpl_scaf_6172-0.3), which can bind caspase and inhibit apoptosis, was up-regulated in bacteriocytes (Fig. 6E). In Drosophila, the baculoviral IAP proteins are important in animal’s response to cellular stress and promote cell survival89,90. Similarly, the gene encoding MAP3K14 (Bpl_scaf_60908-4.14) and potential E3 ubiquitin ligase RNF213 (Bpl_scaf_25983-1.26, Bpl_scaf_1894-1.45, Supplementary Fig. S10), which were apoptosis suppressor or regulator, were up-regulated91. These results suggest that starved bacteriocytes were carrying out cell-type-specific adjustments to cope with stresses. The E3 ubiquitin ligases could also work as intracellular immune sensors of bacterial lipopolysaccharides92. Thus, the encoded protein may be able to active downstream immunological toolkit to digest the symbiont population for nutrients.

Revised version: In starved G. platifrons, negative regulators of cell proliferation, such as autocrine proliferation repressors and receptor-type tyrosine phosphatase beta, were up-regulated in bacteriocytes, suggesting repression of cell growth. We also observed the enrichment of genes in the apoptosis pathway (Supplementary Fig. S8B), which is attributed to the upregulation of caspases (Bpl_scaf_25165-3.8 and Bpl_scaf_24225-0.12) and cathepsins (Bpl_scaf_64706-0.16, Bpl_scaf_61711-5.12) which can trigger caspases-dependent cell death, suggesting cellular stress condition. Correspondingly, a gene encoding baculoviral IAP repeat-containing protein (Bpl_scaf_6172-0.3), which can bind caspase and inhibit apoptosis, was up-regulated in bacteriocytes (Fig. 6E). In Drosophila, the baculoviral IAP proteins are important in animal’s response to cellular stress and promote cell survival89,90. Similarly, the gene encoding MAP3K14 (Bpl_scaf_60908-4.14) and potential E3 ubiquitin ligase RNF213 (Bpl_scaf_25983-1.26, Bpl_scaf_1894-1.45, Supplementary Fig. S10), which were apoptosis suppressor or regulator, were up-regulated91. These results suggest that starved bacteriocytes were carrying out cell-type-specific adjustments to cope with stresses. The E3 ubiquitin ligases could also work as intracellular immune sensors of bacterial lipopolysaccharides92. Thus, the encoded protein may be able to active downstream immunological toolkit to digest the symbiont population for nutrients.

As mentioned above, bacteriocytes obtain nutrients from endosymbionts. KEGG pathway analyses suggested that phagosome, lysosome-related pathways were upregulated in the starvation state, indicating that bacteriocytes were actively digesting the endosymbionts in the starved G. platifrons. Interestingly, protein synthesis activity was more activated in inter lamina cells and VEPCs as supported by the up-regulation of genes encoding ribosomal proteins in both cell populations (Supplementary Tables S6 and S17). This was contrary to the situation in bacteriocytes and may be a consequence of the high activity of “farming” in bacteriocytes.

We anticipated that after moving the starved G. platifrons back to Fanmao site, bacteriocytes and their endosymbionts would “reconstitute” leading to the partly restored “farming” and “milking” pathways. This was confirmed by higher expression of fatty acid metabolic genes in the mussels in the reconstitution state than those in the starvation state, such as long-chain-fatty-acid-ligase ACSBG2-like (Bpl_scaf_28862-1.5), elongation of very long-chain fatty acids 7-like (Bpl_scaf_5959-1.9), and fatty acid desaturase 1-like (Bpl_scaf_35916-0.6). These findings were also consistent with the result of KEGG analyses. The mitochondrial trifunctional enzyme (Bpl_scaf_42376-0.5) was also highly expressed, suggesting a high level of energy-producing activity. The KEGG analyses showed that glutamatergic synthase was enriched in the reconstitution state in comparison with those in both the Fanmao and starvation states. In the insect aphid–Buchnera endosymbiosis model, the host-produced glutamate could be transported and directly utilised by the symbiont93. Similar glutamate-based host–symbiont metabolic interaction mechanisms were also proposed in the deep-sea mussel Bathymodiolus thermophilus and B. azoricus64,94. As the pseudotime analyses showed above (Fig. 6C), reconstitution was the intermediate state between Fanmao and starvation. Metabolic and gene regulatory functions of bacteriocytes were still different from that in Fanmao. The KEGG analyses suggested regulatory (Rap1 signalling, retrograde endocannabinoid signalling, chemokine signalling, glucagon signalling, thyroid hormone synthesis, mRNA surveillance pathway, etc.) and metabolic pathways were enriched in the reconstitution state (fatty acid metabolism, salivary secretion, aldosterone synthesis and section, insulin secretion and pancreatic section) (Supplementary Fig. S8C). For example, we detected up-regulation of the genes encoding carbonic anhydrases (Bpl_scaf_33596-7.3 and Bpl_scaf_48274-0.3), electroneutral sodium bicarbonate exchanger 1 (Bpl_scaf_61230-5.11), and globin-like proteins (Bpl_scaf_50392-1.14 and Bpl_scaf_24370-0.5), which could provide the symbiont with carbon dioxide and oxygen necessary for the symbiont’s chemosynthetic metabolism84,95.

Summary and outlook

Using the deep-sea mussel G. platifrons as a model organism, we demonstrated the power of integrating snRNA-seq and WISH data in unravelling the mechanism behind animal-microbe symbiosis. The robustness of this strategy showed stable and highly distinguishable expression patterns of each cell type regardless of the different environmental states. We successfully profiled the specific roles of different types of cells, including the previously unknown cell types, in maintaining the structure and function of the gill. The supportive cells that located in between (inter lamina cells) and on the basal membrane’s surface (BMC1 and BMC2) helped maintain the anatomical structure of the basal membrane. Ciliary and smooth muscle cells involved in gill slice contraction and cilium beat allowing the gathering of material and information from the surrounding environment. Proliferation cells (PBEZCs, DEPCs and VEPCs) gave rise to new cells, including bacteriocytes which obtained nutrients from endosymbionts using intracellular vacuoles and lysosomes.

The snRNA analyses also revealed that different cell types collaborated to support bacteriocytes’ functionality. The PBEZCs gave rise to new bacteriocytes that allowed symbiont colonisation. Bacteriocytes attached to the basal membrane and were stabilised by supportive cells. Mucus cells co-localise with bacteriocytes and help to maintain immune homeostasis. The beating ciliary cells controlled the water flow, providing bacteriocytes with necessary inorganic substances from the environment. This new information on cell-cell interaction certainly advanced our overall understanding of how endosymbiotic microbes and host cells communicate and collaborate, which cannot be easily achieved through other methods.

Earlier version: Moreover, the analysis of snRNA data from bacteriocytes has provided insight into the molecular mechanisms employed by the host to maintain and regulate the symbiosis. Notably, the bacteriocytes-enriched transcripts involved in harbouring, digesting symbionts, and transporting nutrients produced by the symbionts were identify and characterised. Our in situ transplant experiments by moving mussels between methane-rich and methane-limited sites also provided clues of cell-type specific responses to environmental change. Under a methane-limited environment, the staved mussels more actively consumed endosymbionts through the “farming” pathway. After being moved back to the methane-rich environment, mussels produced more glutamates which sustained the regrowth of symbionts. These preliminary findings showed that the deep-sea mussels were able to control their endosymbionts using a set of genes in response to environmental change. Due to the limitation of remotely operated vehicle (ROV) dives and sampling capacity, we had to have pooled samples of each state for nuclei extraction and sequencing. Thus, the cells per cluster could be considered as technical than biological replicates. Although this sampling process strategy has been broadly used in snRNA/scRNA sequencing96, we recognise the possible violation of assumptions in p-value calculation for DEGs between the three states.

Revised version: Moreover, the analysis of snRNA data from bacteriocytes has provided insight into the molecular mechanisms employed by the host to maintain and regulate the symbiosis. Notably, the bacteriocytes-enriched transcripts involved in harbouring, digesting symbionts, and transporting nutrients produced by the symbionts were identified and characterised. Our in situ transplant experiments by moving mussels between methane-rich and methane-limited sites also provided clues of cell-type specific responses to environmental change. Under a methane-limited environment, the staved mussels more actively consumed endosymbionts through the “farming” pathway. After being moved back to the methane-rich environment, mussels produced more glutamates which sustained the regrowth of symbionts. These preliminary findings showed that the deep-sea mussels were able to control their endosymbionts using a set of genes in response to environmental change. Due to the limitation of remotely operated vehicle (ROV) dives and sampling capacity, we had to have pooled samples of each state for nuclei extraction and sequencing. Thus, the cells per cluster could be considered as technical than biological replicates. Although this sampling process strategy has been broadly used in snRNA/scRNA sequencing96, we recognise the possible violation of assumptions in p-value calculation for DEGs between the three states.

Overall, the single-cell spatial and functional atlas developed in the present work will decipher some common principles of symbiosis and environmental adaption mechanisms of animals. The workflow developed in the present study could provide insightful references for researchers focusing on the mechanistic study of the biological adaptation of biologically and ecologically important non-model animals.

Supporting information